# Statystyka matematyczna - ćwiczenia laboratoryjne

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

# Zestaw zadań 3: Rozkłady zmiennych losowych wielowymiarowych

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Funkcje do wykorzystania

```python
def vizualize_2d_discrete_distribution(pmf, cdf, mask=None):
    """
    Parameters
    -------
    pmf: pandas.DataFrame
        Tablica zawierająca nazwy kolumn z wariantami pierwszej zmiennej losowej i nadaną zbiorczą nazwą zmiennej (np. x),
        indeks z wariantami drugiej zmiennej losowej i nadaną zbiorczą nazwą zmiennej (np. y),
        a wartości w tablicy są prawdopodobieństwami poszczególnych wariantów zmiennej.
    cdf: pandas.DataFrame
        Tablica zawierająca nazwy kolumny w formacie pandas.IntervalIndex z przedziałami wartości pierwszej zmiennej losowej i nadaną zbiorczą nazwą zmiennej (np. x),
        indeks w formacie pandas.IntervalIndex z przedziałami wartości drugiej zmiennej losowej i nadaną zbiorczą nazwą zmiennej (np. y),
        a wartości w tablicy są wartościami dystrybuanty dla poszczególnych przedziałów.
    mask: numpy.ndarray
        Macierz logiczna 2D określająca, które z wariantów zmiennej losowej spełniają wizualizowaną nierówność, a które nie.
    """    
    fig, axes = plt.subplots(1, 2, figsize=(14, 6), facecolor='white')

    fig.suptitle("Rozkład prawdopodobieństwa i dystrybuanta dwuwymiarowej zmiennej losowej dyskretnej")

    if mask is None:
        sns.heatmap(data=pmf, vmin=0, vmax=np.max(pmf.values), cmap="Blues", annot=True, linewidths=2, linecolor='grey', ax=axes[0])
    else:
        sns.heatmap(data=pmf, vmin=0, vmax=np.max(pmf.values), cmap="Blues", annot=True, annot_kws={'color':'darkgreen'}, linewidths=2, linecolor='grey', mask=~mask, ax=axes[0])
        sns.heatmap(data=pmf, vmin=0, vmax=np.max(pmf.values), cmap="Blues", annot=True, annot_kws={'color':'red'}, linewidths=2, linecolor='grey', mask=mask, cbar=False, ax=axes[0])
    axes[0].set_title("Rozkład prawdopodobieństwa")
    axes[0].tick_params(labelrotation=0)

    sns.heatmap(data=cdf, vmin=0, vmax=1, cmap="Blues", annot=True, linewidths=2, linecolor='grey', ax=axes[1])
    axes[1].set_title("Dystrybuanta")
    axes[1].tick_params(labelrotation=45)

```

### Dane do zadań

Pliki r1p.pkl oraz r1f.pkl zawierają funkcję rozkładu prawdopodobieństwa oraz dystrybuantę przykładowej dwywymiarowej zmiennej losowej dyskretnej.

W celu wczytania danych do notatnika umieść pliki w tym samym folderze, w którym znajduje się notatnik.

```python
r1p = pd.read_pickle("r1p.pkl")
r1p
```

Output:
```
File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\io\pickle.py:185, in read_pickle(filepath_or_buffer, compression, storage_options)
    123 """
    124 Load pickled pandas object (or any object) from file.
    125 
   (...)    182 4    4    9
    183 """
    184 excs_to_catch = (AttributeError, ImportError, ModuleNotFoundError, TypeError)
--> 185 with get_handle(
    186     filepath_or_buffer,
    187     "rb",
    188     compression=compression,
    189     is_text=False,
    190     storage_options=storage_options,
    191 ) as handles:
    192     # 1) try standard library Pickle
    193     # 2) try pickle_compat (older pandas version) to handle subclass changes
    194     # 3) try pickle_compat with latin-1 encoding upon a UnicodeDecodeError
    196     try:
    197         # TypeError for Cython complaints about object.__new__ vs Tick.__new__
    198         try:

File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\io\common.py:882, i
```

```python
r1f = pd.read_pickle("r1f.pkl")
r1f
```

```python
vizualize_2d_discrete_distribution(r1p, r1f)
```

### Zadanie 1: Przykładowy dwuwymiarowy dyskretny rozkład prawdopodobieństwa [17 pkt]

Bazując na wczytanym do notatnika przykładowym dwuwymiarowym rozkładzie prawdopodobieństwa oblicz jego rozkłady brzegowe, wybrane rozkłady warunkowe oraz podstawowe parametry rozkładu.

#### a) Rozkłady brzegowe

Przedstaw funkcje brzegowych rozkładów prawdopodobieństwa zmiennej losowej dwuwymiarowej X, Y w formie tablic `r1bx`  i `r1by` o następującej strukturze:

Tablica `r1bx`:
> Indeks:</br>
> x - warianty przyjmowane przez zmienną X.</br>

> Kolumny: </br>
> P(X=x) - prawdopodobieństwa rozkładu brzegowego zmiennej X.

Tablica `r1by`:
> Indeks:</br>
> y - warianty przyjmowane przez zmienną Y.</br>

> Kolumny: </br>
> P(Y=y) - prawdopodobieństwa rozkładu brzegowego zmiennej Y.

<!-- grade_id: cell-9c3fcdd9befa96eb, points: 0, grade: false, solution: true -->
```python
# Rozkład brzegowy X
r1bx = pd.DataFrame(
    data = np.sum(r1p,axis=0).values,
    index = r1p.columns,
    columns = ["P(X=x)"]
)
r1bx.index.name="x"
r1bx
```

<!-- grade_id: cell-60e0b6c63085aa73, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1bx) == pd.DataFrame
assert r1bx.shape == (4, 1)
assert list(r1bx.columns) == ["P(X=x)"]
assert r1bx.index.name == 'x'
assert np.isclose(np.sum(r1bx.index.values), 10)
assert np.isclose(np.sum(r1bx["P(X=x)"]), 1)
```

<!-- grade_id: cell-47892605f8f5f38f, points: 0, grade: false, solution: true -->
```python
# Rozkład brzegowy Y
r1by = pd.DataFrame(
    data = np.sum(r1p,axis=1).values,
    index = r1p.index.values,
    columns = ["P(Y=y)"]
)
r1by.index.name="y"
r1by
```

<!-- grade_id: cell-33fbf4e898073a05, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1by) == pd.DataFrame
assert r1by.shape == (4, 1)
assert list(r1by.columns) == ["P(Y=y)"]
assert r1by.index.name == 'y'
assert np.isclose(np.sum(r1by.index.values), 30)
assert np.isclose(np.sum(r1by["P(Y=y)"]), 1)
```

#### b) Przykładowe rozkłady warunkowe

Przedstaw funkcje przykładowych warunkowych rozkładów prawdopodobieństwa zmiennej losowej dwuwymiarowej X, Y w formie tablic `r1wx`  i `r1wy` o następującej strukturze:

Tablica `r1wx`:
> Indeks:</br>
> x - warianty przyjmowane przez zmienną X.</br>

> Kolumny: </br>
> P(X=x|Y=5) - prawdopodobieństwa rozkładu warunkowego P(X=x|Y=5) zmiennej X.

Tablica `r1wy`:
> Indeks:</br>
> y- warianty przyjmowane przez zmienną Y.</br>

> Kolumny: </br>
> P(Y=y|X=2) - prawdopodobieństwa rozkładu warunkowego P(Y=y|X=2) zmiennej Y.

```python

```

<!-- grade_id: cell-4f06b6c68b551d3c, points: 0, grade: false, solution: true -->
```python
# Rozkład warunkowy X

# P_B = r1by['P(Y=y)'][5]
# for i in r1p.columns.values:
    # print(r1p[i][5] / P_B)


# loc bardzo przydatny, zamiast robienia forow i dziwniejszych odnosien do tablic
# kacper zrobił trochę inaczej, wykorzystując np. loc[5,:]

r1wx = pd.DataFrame(
    data = (r1p.loc[5]).values / r1by.loc[5].values,
    index = r1p.columns,
    columns = ["P(X=x|Y=5)"]
)
r1wx.index.name="x"
r1wx
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 12
      1 # Rozkład warunkowy X
      2 
      3 # P_B = r1by['P(Y=y)'][5]
   (...)      8 # loc bardzo przydatny, zamiast robienia forow i dziwniejszych odnosien do tablic
      9 # kacper zrobił trochę inaczej, wykorzystując np. loc[5,:]
     11 r1wx = pd.DataFrame(
---> 12     data = (r1p.loc[5]).values / r1by.loc[5].values,
     13     index = r1p.columns,
     14     columns = ["P(X=x|Y=5)"]
     15 )
     16 r1wx.index.name="x"
     17 r1wx

NameError: name 'r1p' is not defined
```

<!-- grade_id: cell-8555855320b32314, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1wx) == pd.DataFrame
assert r1wx.shape == (4, 1)
assert list(r1wx.columns) == ["P(X=x|Y=5)"]
assert r1wx.index.name == 'x'
assert np.isclose(np.sum(r1wx.index.values), 10)
assert np.isclose(np.sum(r1wx["P(X=x|Y=5)"]), 1)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 2
      1 # Komórka testowa
----> 2 assert type(r1wx) == pd.DataFrame
      3 assert r1wx.shape == (4, 1)
      4 assert list(r1wx.columns) == ["P(X=x|Y=5)"]

NameError: name 'r1wx' is not defined
```

<!-- grade_id: cell-659e50116160d158, points: 0, grade: false, solution: true -->
```python
# Rozkład warunkowy Y

# (r1p.loc[5]).values / r1by.loc[5].values

# P_B 

r1p

# r1wy = pd.DataFrame(
#     data = ,
#     index = r1p.index,
#     columns = ["P(Y=y|X=2)"]
# )
# r1wy.index.name="y"
# r1wy
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[1], line 7
      1 # Rozkład warunkowy Y
      2 
      3 # (r1p.loc[5]).values / r1by.loc[5].values
      4 
      5 # P_B 
----> 7 r1p
      9 # r1wy = pd.DataFrame(
     10 #     data = ,
     11 #     index = r1p.index,
   (...)     14 # r1wy.index.name="y"
     15 # r1wy

NameError: name 'r1p' is not defined
```

<!-- grade_id: cell-c0434f21c515cbf3, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1wy) == pd.DataFrame
assert r1wy.shape == (4, 1)
assert list(r1wy.columns) == ["P(Y=y|X=2)"]
assert r1wy.index.name == 'y'
assert np.isclose(np.sum(r1wy.index.values), 30)
assert np.isclose(np.sum(r1wy["P(Y=y|X=2)"]), 1)
```

#### c) Podstawowe parametry rozkładu

Oblicz i zapisz do zmiennych następujące parametry rozkładu zmiennej losowej dwuwymiarowej X, Y - wartość średnia (`r1_mean`), wariancja(`r1_var`), odchylenie standardowe (`r1_std`), kowariancję (`r1_cov`) oraz współczynnik korelacji (`r1_r`).

<!-- grade_id: cell-6549962e1f1f7081, points: 0, grade: false, solution: true -->
```python
# Wartość oczekiwana
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0edb5536805a6a81, points: 0, grade: false, solution: true -->
```python
# Wariancja
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5159094f21b704e7, points: 0, grade: false, solution: true -->
```python
# Odchylenie standardowe
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43093815f03a1550, points: 0, grade: false, solution: true -->
```python
# Kowariancja
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b87e1ff371fb5862, points: 0, grade: false, solution: true -->
```python
# Współczynnik korelacji
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bf0094e8618c1359, points: 9, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(r1_mean, 18.505)
assert np.isclose(r1_var, 38.96800412)
assert np.isclose(r1_std, 6.242435752172384)
assert np.isclose(r1_cov, 0.07360000000000022)
assert np.isclose(r1_r, 0.011848900761810917)
```

### Zadanie 2: Obliczanie i wizualizacja prawdopodobieństw [6 pkt]

Oblicz i przedstaw na wykresach następujące prawdopodobieństwa:
- P(X>2, 0<Y<=10)
- P(X<=2|X>3, Y>=5)

Wyniki zapisz kolejno do zmiennych `p1` i `p2`oraz przedstaw na wykresach korzystając z funkcji `vizualize_2d_discrete_distribution()`.

#### a) P(X>2, 0<Y<=10)

<!-- grade_id: cell-902a8df721d4260e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-be46ed7911a5be86, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(p1, 0.17)
```

#### b) P(X<=2|X>3, Y>=5)

<!-- grade_id: cell-1478dbd04b61d3fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-572d8bc3d9228e82, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(p2, 0.362)
```

### Zadanie 3: Niezależność zmiennych losowych [2 pkt]

Sprawdź, czy zmienne losowe X i Y są zmiennymi niezależnymi i zapisz wynik testu (True/False) do zmiennej `test`.

<!-- grade_id: cell-bb7d573597e9f308, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0b3ec359911df3b8, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert test==False
```