# Statystyka matematyczna - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autora rozwiązania (imię, nazwisko i numer indeksu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Imię i nazwisko:
Numer indeksu:

---

# Kolowkium zaliczeniowe: 1 termin, grupa 7

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

## Dane do automatycznych testów

```python
test_data_1 = pd.DataFrame(data=np.concatenate([sp.stats.norm.rvs(loc=4, scale=0.2, size=45, random_state=7), sp.stats.norm.rvs(loc=5.5, scale=0.4, size=85, random_state=2159)]), columns=["X"])
test_data_1.head()
```

```python
test_data_2 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=15, random_state=7), columns=["X"])
test_data_2.head()
```

## Zadanie 1 [4 pkt]

Na pewnym skrzyżowaniu dochodzi średnio do 2 stłuczek dziennie.

Jakie jest prawdopodobieństwo, że w ciągu najbliższego tygodnia (7 dni) będzie miało miejsce co najmniej 15 stłuczek?

Wynik zapisz do zmiennej `p`.

<!-- grade_id: cell-328b607aafafa4ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e646a3ab55f95f64, points: 4, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(p, float)

```

## Zadanie 2 [8 pkt]

Przygotuj funkcję `simple_kde_estimation()`, która będzie dokonywała estymacji nieparametrycznej metodą funkcji jądrowych z wykorzystaniem jądra w postaci rozkładu normalnego o odchyleniu normalnym $\sigma$ zgodnie z informacjami zawartymi w docstring funkcji oraz schematem z załączonego do notatnika zestawu wzorów.

Poza zwracaniem danych określonych w docstring funkcji powinna ona generować wykres zawierający wyestymowaną funkcję (wykres liniowy z wypełnionym polem pod krzywą) oraz dodaną wizualną reprezentacją zestawu danych w postaci zbioru symboli (| lub +) zlokalizowanych poniżej wykresu.

Uwagi do wykonania zadania:
 - Celem zadania jest napisanie funkcji od podstaw, w rozwiązaniu nie należy korzystać z gotowych funkcji do estymacji KDE.
 - Przygotowując figurę kieruj się wytycznymi z początku notatnika.

<!-- grade_id: cell-00b25bd896944c4d, points: 0, grade: false, solution: true -->
```python
def simple_kde_estimation(data, estimation_points, sigma):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i pojedynczą kolumnę "X" z wartościami próby losowej.
    estimation_points: int
        Liczba punktów estymacji funkcji gęstości prawdopodobieństwa.
    sigma: float
        Odchylenie standardowe rozkładu normalnego wykorzystywanego jako jądro funkcji.
    
    Returns
    -------
    pdf: DataFrame
        Tablica zawierająca następujące kolumny:
        X - uszeregowane rosnąco punkty estymacji funkcji gęstości prawdopodobieństwa,
        KDE - wyestymowane wartości funkcji gęstości prawdopodobieństwa.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-60ea6c2bcb228b58, points: 8, grade: true, solution: false -->
```python
# Komórka testowa
test_results_1 = simple_kde_estimation(test_data_1, estimation_points=100, sigma=0.2)
test_results_2 = simple_kde_estimation(test_data_1, estimation_points=250, sigma=0.35)

assert type(test_results_1) == pd.DataFrame
assert type(test_results_2) == pd.DataFrame

```

## Zadanie 3 [6 pkt]

Przygotuj funkcję `variance_NHST()`, która będzie weryfikować hipotezę o wartości wariancji dla danych wejściowych w postaci szeregu szczegółowego zgodnie z informacjami zawartymi w docstring funkcji oraz ze schematem z załączonego do notatnika zestawu wzorów.

Uwaga! Zadanie dotyczy tradycyjnej weryfikacji, nie metody bootstrap.

<!-- grade_id: cell-146cf426bb90e71e, points: 0, grade: false, solution: true -->
```python
def variance_NHST(data, H0, H1, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i kolumnę "X" wynikami próby losowej.
    H0: float
        Wariancji przyjęta jako hipoteza zerowa.
    H1: str
        Postać hipotezy alternatywnej, przyjmuje wartości:
        - two-sided: wariancja jest różna od wartości przyjętej w H0,
        - less: wariancja jest mniejsza od wartości przyjętej w H0,
        - greater: wariancja jest większa od wartości przyjętej w H0.
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    chi2: float
        Wyliczona na podstawie próby losowej wartość statystyki chi2.
    chi2_alpha: float or list
        Wartość statystyki chi2_alpha lub lista z dwoma wartościami chi2_alpha_1 i chi2_alpha_2
    H: int
        Wynik testu statystycznego, przyjmuje wartość:
        0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha,
        1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aa9e98593265cfc8, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
result = variance_NHST(test_data_2, 0.1, "two-sided", 0.05)
assert type(result)==tuple or type(result)==list
assert len(result)==3
assert isinstance(result[0], float)
assert type(result)==tuple or type(result)==list
assert isinstance(result[2], float) or isinstance(result[2], int)

result = variance_NHST(test_data_2, 0.1, "less", 0.05)
assert type(result)==tuple or type(result)==list
assert len(result)==3
assert isinstance(result[0], float)
assert isinstance(result[0], float)
assert isinstance(result[2], float) or isinstance(result[2], int)

```

## Zadanie 4 [10 pkt]

Zmienna `penguins` zawiera wyniki badania pingwinów przeprowadzonych na wyspach Archipelagu Palmera (Antarktyda):
 - species - gatunek pingwina
 - culmen_length_mm - długość dzioba (mm)
 - culmen_depth_mm - głębokość dzioba (mm)
 - flipper_length_mm - długość płetwy (mm)
 - body_mass_g - masa ciała (g)
 - island - nazwa wyspy
 - sex - płeć pingwina

```python
penguins = pd.read_csv("penguins_data.csv", index_col=0)
penguins
```

### a)
W oparciu o dane ze zmiennej `penguins` przygotuj wykresy violinplot i stripplot wybranego przez siebie parametru liczbowego (oś Y) z podziałem na rozróżnione kolorem gatunki pingwinów (oś X).

Wykresy zestaw na pojedynczej figurze (1x2 pola na wykres). Przygotowując figurę kieruj się wytycznymi z początku notatnika. Dodatkowo ustaw:
 - dla wykresu violinplot:
     - inne od domyślnych kolory wykresu,
 - dla wykresu stripplot:
     - inne od domyślnych rozmiar, kolory i przeźroczystość punktów (dobierz parametry tak, aby wykres był jak najbardziej przejrzysty) oraz rozsunięcie poszczególnych grup punktów.

<!-- grade_id: cell-da00623fe210a9ed, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### c)

Załóżmy, że zmianna `penguins` zawiera pełną populację pingwinów zamieszkujących poszczególne wyspy. Wylosuj z niej i zapisz do zmiennej `sample` próbę losową zawierającą po 20% losowo wybranych osobników zamieszkujących każdą z wysp (bez zwracania).

<!-- grade_id: cell-6060bc71a588c481, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eced5f1fe29d8b10, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(sample) == pd.DataFrame

```

### c)

Korzystając z dedykowanej do tego funkcji pakietu scipy dokonaj estymacji przedziałowej metodą bootstrap współczynnika a prostej regresji liniowej (ax+b) między długością płetw a masą ciała pingwinów na poziomie istotności 1-alpha = 0.98.

Wynik zapisz do zmiennej `a_ci` zawierającej listę lub krotkę z dolną i górną granicą przedziału ufności. Ustaw wartość argumentu funkcji odpowiadającego za ziarno generatora liczb losowych funkcji na 34.

<!-- grade_id: cell-b6f52487824e8671, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-94f8e37e60f83338, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(a_ci)==tuple or type(a_ci)==list
assert len(a_ci)==2
assert isinstance(a_ci[0], float)
assert isinstance(a_ci[1], float)

```

### d)

Korzystając z dedykowanej funkcji pakietu scipy zweryfikuj hipotezę, że średnia długość dzioba pingwinów z gatunku Gentoo jest równa 48 mm względem hipotezy alternatywnej, że jest mniejsza od 48 mm. W obliczeniach przyjmij poziom istotności 1 - alpha = 0.98.

Zapisz obliczoną wartość pvalue do zmiennej `pvalue`, oraz przygotuj zmienną `H` zawierającą wynik testu statystycznego (0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha, 1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha).

<!-- grade_id: cell-b644105f89af5513, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2dd896b4ca1ac760, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(pvalue, float)
assert isinstance(H, float) or isinstance(H, int)

```