# Statystyka matematyczna - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autora rozwiązania (imię, nazwisko i numer indeksu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Imię i nazwisko:
Numer indeksu:

---

# Kolowkium zaliczeniowe: 1 termin, grupa 4

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

## Dane do automatycznych testów

```python
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=25, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=25, random_state=13)
test_data_1 = pd.DataFrame(data=np.array([x, y]).T, columns=["X", "Y"])
test_data_1.head()
```

Output:
```
X         Y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

```python
test_data_2 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=15, random_state=7), columns=["X"])
test_data_2.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

## Zadanie 1 [4 pkt]

Na pewnym skrzyżowaniu dochodzi średnio do 2 stłuczek dziennie.

Jakie jest prawdopodobieństwo, że w ciągu najbliższego tygodnia (7 dni) będzie miało miejsce co najmniej 15 stłuczek?

Wynik zapisz do zmiennej `p`.

<!-- grade_id: cell-328b607aafafa4ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e646a3ab55f95f64, points: 4, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(p, float)

```

## Zadanie 2 [6 pkt]

Przygotuj funkcję `bootstrap_correlation_estimation()`, która korzystając z metody bootstrap będzie dokonywała estymacji przedziałowej wartości współczynnika korelacji liniowej Pearsona dla danych wejściowych w postaci szeregu szczegółowego zgodnie z informacjami zawartymi w docstring funkcji oraz ze schematem z załączonego do notatnika zestawu wzorów.

Uwagi do wykonania zadania:
 - Wykonaj losowanie prób wg następującego schematu - kolejne próby są losowane w pęli for, a funkcja losująca próbę w kolejnych losowaniach ma ustawioną wartość parametru odpowiadającego za ziarno generatora liczb losowych na wartość `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwiązaniu nie należy korzystać z gotowych funkcji służących do estymacji bootstrapowej.

<!-- grade_id: cell-00b25bd896944c4d, points: 0, grade: false, solution: true -->
```python
def bootstrap_correlation_estimation(data, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i dwie kolumny "X" i "Y" z wynikami próby losowej.
    number_of_samples: int
        Liczba prób bootstrapowych.
    alpha: float
        Wartość określająca poziom istotności.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziału ufności.
    r_corr_high: float
        Górna granica wyliczonego przedziału ufności.
    """
    wyniki_bootstrap = []

    for i in range(number_of_samples):
        probka_bootstrap = data.sample(n=len(data), replace=True, random_state=random_state + i)
        korelacja_pearsona = probka_bootstrap['X'].corr(probka_bootstrap['Y'])
        wyniki_bootstrap.append(korelacja_pearsona)
    
    dolna_granica = alpha / 2 * 100
    gorna_granica = (1 - alpha / 2) * 100
    r_corr_low = np.percentile(wyniki_bootstrap, dolna_granica)
    r_corr_high = np.percentile(wyniki_bootstrap, gorna_granica)

    poziom_ufnosci = 1 - alpha
    print(f'r należy do przedziału [{r_corr_low:.4f}, {r_corr_high:.4f}] przy założeniu poziomu ufności 1-α = {poziom_ufnosci:.2f}')
    
    return r_corr_low, r_corr_high
```

<!-- grade_id: cell-60ea6c2bcb228b58, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
result = bootstrap_correlation_estimation(test_data_1, 1200, 0.05, 14)
assert type(result)==tuple or type(result)==list
assert len(result)==2
assert isinstance(result[0], float)
assert isinstance(result[1], float)

```

Output:
```
r należy do przedziału [0.9291, 0.9818] przy założeniu poziomu ufności 1-α = 0.95
```

## Zadanie 3 [6 pkt]

Przygotuj funkcję `mean_NHST()`, która będzie weryfikować hipotezę o wartości przeciętnej dla danych wejściowych w postaci szeregu szczegółowego zgodnie z informacjami zawartymi w docstring funkcji oraz ze schematem z załączonego do notatnika zestawu wzorów.

Uwaga! Zadanie dotyczy tradycyjnej weryfikacji, nie metody bootstrap.

<!-- grade_id: cell-146cf426bb90e71e, points: 0, grade: false, solution: true -->
```python
def mean_NHST(data, H0, H1, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i kolumnę "X" wynikami próby losowej.
    H0: float
        Wartość przeciętna przyjęta jako hipoteza zerowa.
    H1: str
        Postać hipotezy alternatywnej, przyjmuje wartości:
        - two-sided: wartość przeciętna jest różna od wartości przyjętej w H0,
        - less: wartość przeciętna jest mniejsza od wartości przyjętej w H0,
        - greater: wartość przeciętna jest większa od wartości przyjętej w H0.
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    t: float
        Wyliczona na podstawie próby losowej wartość statystyki t.
    t_alpha: float
        Wartość statystyki t_alpha.
    H: int
        Wynik testu statystycznego, przyjmuje wartość:
        0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha,
        1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha.
    """
    x_bar = data["X"].mean() # średnia z danych
    s = data["X"].std(ddof=1) #standardowe odchylenie nie działa bez ddof(nieobciążony estymator)
    n = len(data)            #ilość 

    #końcowy wzór etap 2: karta wzorów 
    t = (x_bar - H0) * np.sqrt(n)/ s

    # Etap 4: Ustalenie poziomu istotności α
   #nasza alpha

    # Etap 5: Wyznaczenie obszaru krytycznego testu istotności
    if H1 == "two-sided":
        t_alpha = stats.t.ppf(1 - alpha / 2,  df=(n - 1))

      #obszar krytyczny dla funkcji wizualizacyjnej

    #  rejection_region: list
      #  Lista zawierająca trzy elementy pozwalające na określenie obszaru krytycznego w następującej kolejności:
      #   - wartość dolnej granicy przedziału,
      #   - wartość górnej granicy przedziału,
      #   - "inner" lub "outer" w zależności od tego, czy ma być zakolorowana wewnętrzna, czy zewnętrzna część przedziału.
  
        rejection_region = [-t_alpha, t_alpha, "outer"]
        
    elif H1 == "less":
        t_alpha = stats.t.ppf(alpha,  df=(n - 1))
        rejection_region = [-np.inf, t_alpha, "inner"]
        
    else:  # H1 == "greater"
        t_alpha = stats.t.ppf(1 - alpha, df=(n - 1))
        rejection_region = [t_alpha, np.inf, "inner"]

    # Etap 6: Podjęcie decyzji weryfikującej
    if H1 == "two-sided":
        if abs(t) >= t_alpha:
            H = 1  # Odrzucenie H0 
        else:
            H = 0  # Brak  odrzucenia H0
    elif H1 == "less":
        if t <= t_alpha:
            H = 1  # Odrzucenie H0 
        else:
            H = 0  # Brak  odrzucenia H0
    else:  # H1 == "greater"
        if t >= t_alpha:
            H = 1  # Odrzucenie H0 
        else:
            H = 0  # Brak odrzucenia H0
    
    return t, t_alpha, H
```

<!-- grade_id: cell-aa9e98593265cfc8, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
result = mean_NHST(test_data_2, 5.08, "two-sided", 0.05)
assert type(result)==tuple or type(result)==list
assert len(result)==3
assert isinstance(result[0], float)
assert isinstance(result[1], float)
assert isinstance(result[2], float) or isinstance(result[2], int)

```

## Zadanie 4 [12 pkt]

Zmienna `hawks` zawiera wyniki badania jastrzębi nad jeziorem MacBride w pobliżu Iowa City:
 - BandNumber - kod opaski identyfikacyjnej
 - Species - gatunek
 - Wing - długość (w mm) głównego pióra skrzydła
 - Hallux - długość (w mm) szponu
 - Culmen - długość (w mm) górnej części dzioba
 - Tail - pomiar (w mm) związany z długością ogona

```python
hawks = pd.read_csv("hawks_data.csv", index_col=0)
hawks
```

Output:
```
BandNumber     Species   Wing  Hallux  Culmen  Tail
0     877-76321  Red-tailed  375.0    30.0    27.2   243
1     877-76322  Red-tailed  412.0    31.3    29.3   210
2     877-76325  Red-tailed  390.0    30.9    28.9   213
3    1207-55917  Red-tailed  393.0    30.6    28.2   238
4    1207-55913  Red-tailed  436.0    34.0    30.5   232
..          ...         ...    ...     ...     ...   ...
211  1177-04773  Red-tailed  370.0    31.4    26.9   201
212  1177-04774  Red-tailed  360.0    30.6    26.2   224
213  1177-04775  Red-tailed  366.0    27.2    25.3   199
214  1177-04776  Red-tailed  402.0    31.0    28.7   219
215  1207-53145  Red-tailed  199.0    32.1    28.7   222

[216 rows x 6 columns]
```

### a)
W oparciu o dane ze zmiennej `hawks` przygotuj zestaw wykresów korelacyjnych poszczególnych par parametrów z wykresami kde na przekątnej (pairplot) zestawiające ze sobą poszczególne parametry liczbowe (osie X, Y) z podziałem na gatunki jastrzębi (rozróżnione kolorem) i ustaw:
 - inną od domyślnej skalę kolorów (każdy gatunek ma przypisany inny kolor wykresu),
 - inne od domyślnych własności punktów (kształt, wielkość, przeźroczystość) dobrane tak, aby wykres był jak najbardziej czytelny.

<!-- grade_id: cell-c8281a216810df8a, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### b)
W oparciu o dane ze zmiennej `hawks` przygotuj wykresy boxplot i stripplot wybranego przez siebie parametru liczbowego (oś Y) z podziałem na rozróżnione kolorem gatunki jastrzębi (oś X).

Wykresy zestaw na pojedynczej figurze (1x2 pola na wykres). Przygotowując figurę kieruj się wytycznymi z początku notatnika. Dodatkowo ustaw:
 - dla wykresu boxplot:
     - inne od domyślnych kolory wykresu oraz widoczny znacznik wartości średniej,
 - dla wykresu stripplot:
     - inne od domyślnych rozmiar, kolory i przeźroczystość punktów (dobierz parametry tak, aby wykres był jak najbardziej przejrzysty) oraz rozsunięcie poszczególnych grup punktów.

<!-- grade_id: cell-da00623fe210a9ed, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### c)

Załóżmy, że zmianna `hawks` zawiera pełną populację jastrzębi zamieszkujących dany obszar. Wylosuj z niej i zapisz do zmiennej `sample` próbę losową zawierającą po 20 losowo wybranych osobników każdego z gatunków (bez zwracania).

<!-- grade_id: cell-da4aa5c84d38fab7, points: 0, grade: false, solution: true -->
```python
sample = hawks.sample(n=20, replace=True)
sample
```

Output:
```
BandNumber        Species   Wing  Hallux  Culmen  Tail
14    877-76337     Red-tailed  378.0   31.30   25.40   204
97   2003-99315  Sharp-Shinned  205.0   15.50   12.25   169
43   1387-79183     Red-tailed  345.0   30.10   26.40   200
41   1387-79176     Red-tailed  350.0   29.10   26.00   205
99   1177-14621     Red-tailed  404.0   31.55   28.20   224
74   1207-72613     Red-tailed  397.0   31.40   28.80   216
160  1207-72662     Red-tailed  404.0   32.80   29.40   224
28    872-09615  Sharp-Shinned  170.0   11.20   10.10   126
138  1177-04677     Red-tailed  385.0   30.90   27.30   216
138  1177-04677     Red-tailed  385.0   30.90   27.30   216
146  2003-99330  Sharp-Shinned  203.0   14.20   13.50   157
96   1207-72624     Red-tailed  404.0   32.50   30.00   220
108   803-05933  Sharp-Shinned  190.0   14.60   12.80   147
61   1387-92154     Red-tailed  382.0   34.70   34.90   247
115  1207-72642     Red-tailed  408.0   32.70   28.80   225
71   1207-72612     Red-tailed  392.0   29.00
```

<!-- grade_id: cell-39dbe6b5f805f78e, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(sample) == pd.DataFrame

```

### d)

Korzystając z dedykowanej do tego funkcji pakietu scipy dokonaj estymacji przedziałowej metodą bootstrap odchylenia standardowego długości piór w skrzydłach jastrzębi gatunku Sharp-Shinned na poziomie istotności 1-alpha = 0.9.

Wynik zapisz do zmiennej `std_ci` zawierającej listę lub krotkę z dolną i górną granicą przedziału ufności. Ustaw wartość argumentu funkcji odpowiadającego za ziarno generatora liczb losowych funkcji na 17.

<!-- grade_id: cell-b6f52487824e8671, points: 0, grade: false, solution: true -->
```python
dane = hawks[hawks["Species"] == "Sharp-Shinned"]["Wing"]
std_ = sp.stats.bootstrap((dane,), np.std, confidence_level=0.9, random_state=17)
std_lower, std_upper = std_.confidence_interval
std_ci = (std_lower, std_upper)
std_ci
```

Output:
```
(15.207541129759127, 17.847028694919924)
```

<!-- grade_id: cell-94f8e37e60f83338, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(std_ci)==tuple or type(std_ci)==list
assert len(std_ci)==2
assert isinstance(std_ci[0], float)
assert isinstance(std_ci[1], float)

```

### e)

Korzystając z dedykowanej funkcji pakietu scipy zweryfikuj za pomocą testu chi-kwadrat hipotezę, że długość szponów jastrzębi gatunku Red-tailed ma rozkład normalny o wartości oczekiwanej 31 i odchyleniu standardowym 2.5. W obliczeniach przyjmij poziom istotności 1 - alpha = 0.95. 

Podczas obliczeń dokonaj podziału danych na następujące przedziały: [22, 23), [23, 24), ... , [35, 36].

Zapisz obliczoną wartość pvalue do zmiennej `pvalue` oraz przygotuj zmienną `H` zawierającą wynik testu statystycznego (0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha, 1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha).

<!-- grade_id: cell-b644105f89af5513, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2dd896b4ca1ac760, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert isinstance(pvalue, float)
assert isinstance(H, float) or isinstance(H, int)

```